```python
# Initialize Otter
import otter
grader = otter.Notebook("HW1_Fall2022.ipynb")
```

Name: An Yan
Student Number:30167924

# DATA 601

## HW1


**Learning Objectives**

- Explore built-in data types in Python.
- Review fundamental programming and problem solving concepts with Python.
- Implement functions based on mathematical concepts and definitions.
- Gain experience working with the Jupyter notebook environment.

_This is an individual homework assignment._ 

Please complete this homework assignment within the Jupypter notebook environment.  

#### Submission 

In order to ensure that everything goes smoothly, and is easy to grade please follow these instructions:

- Please provide your solutions where asked; please do not alter any other parts of this notebook.


#### <font color='red'> Note: Passing the provided test-cases doesn't guarantee points. We check your code with other test-cases for grading.</font> 

```python
# Check that we are using a recent version of Jupyter.
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Part A

This part focuses on scalar types. You should be able to complete the following questions without using any collection types.

### Question 1

The central bionomial coefficients form a vertical row down the middle of Pascal's triangle. They can be calculated via the following formula:

$$
{2n\choose n} = \frac{(2n)!}{(n!)^2} = \prod_{k=1}^{n} \frac{n+k}{k}  
$$

Write a function that calculates the nth central bionomial coefficient. You may use in-built functions to calculate the factorials.

(3 points)

```python
import math
def cent_binomial_coef(n):
    '''Returns the n-th binomial coefficient for an input
    integer n.
    '''
    num=0
    if n==0:
        return 1
    else:
        for i in range(0, n):
            num=(math.factorial(2*n))/((math.factorial(n))**2)
    return num
    ...
```

```python
grader.check("q1_a")
```

Output:
```
q1_a results: All test cases passed!
```

### Question 2

Write a function that determines whether an integer is prime. You may not use built-in functions to determine this directly.

(8 points)

```python
def is_prime(n):
    '''For a positive integer n, returns True if n is a 
    prime number, False otherwise.
    '''
    for i in range(2, n):
        if n%i==0:
            return False
    return True
    ...
```

```python
grader.check("q2_a")
```

Output:
```
q2_a results: All test cases passed!
```

##### Question 3

(5 points)

Many things in nature can be modeled by the exponential growth function: 
$$
P(t) = \alpha \beta^{\frac{t}{\tau}},
$$
where $\alpha$ is the initial population, $\beta$ is the growth rate factor, t is the time passed, and $\tau$ is the time constant. 

Imagine that you are an entomologist on the plant Zergo. There is a small insect of interest to you on this planet which has an exponential growth rate. From experimentation you determine that: 
$\beta= \sqrt\frac{5}{4}$,  and $\tau= 2^{\frac{1}{3}}$.

Starting with a population $\alpha$ of 10, use this formula to compute $P(t)$ in the function `population(t)` below. 

<!--
BEGIN QUESTION
name: A3
manual: true
points: 2
-->

```python
import math
def population(t):
    '''Computes the population at time t, where t is a positive float, 
    and returns the population as an integer (floor value). 
    It is assumed that time is in minutes and that all parameters
    are correct for this assumption.
    '''
    alpha=10
    beta=math.sqrt(5/4)
    tao=2**(1/3)
    result=alpha*(beta**(t/tao))
    return math.floor(result)
    ...
```

```python
grader.check("q3_a")
```

Output:
```
q3_a results: All test cases passed!
```

## Part B

Questions in this part rely on strings, lists, and tuples.

### Question 1

Consider the lists $a=[2,1,2,5,3,1]$ and $b=[2,1,7,1]$. The intersection of these lists, $a \cap b$, is $[1,1,2]$. Write a function that returns the intersection of two arbitrary lists. You may not use built-in set operations, and should pay attention to duplicates.

(5 points)


```python
def intersection_list(a,b):
    '''
    Returns the intersection of two lists a and b. 
    
    '''
    res=[]
    for i in a:
        if i in b:
            res.append(i)
            b.remove(i)
    return res
    ...
```

```python
grader.check("q1_b")
```

Output:
```
q1_b results: All test cases passed!
```

### Question 2
(8 points)

The scalar triple product of three vectors is defined as:

$$
a.(b \times c)
$$

Write a function that calculates the scalar triple product of three vectors, a, b, and c.

```python
def striple_prod(a,b,c):
    '''Returns the scalar triple product of three vectors A, B, and C. 
    '''
    corss=[]
    res=0
    if len(a)!=len(b) or len(b)!=len(c) or len(a)!=len(c):
        print("Wrong szie")
    else:
        cross=[b[0]*c[2]-c[1]*b[2], b[2]*c[0]-b[0]*c[2], b[0]*c[1]-b[1]*c[0]]
        for i, j in zip(a, cross):
            res=res+i*j
    return res
    ...
```

```python
grader.check("q2_b")
```

Output:
```
q2_b results: All test cases passed!
```

### Question 3
(10 points)

Write a function that, given a piece of text, calculates the most frequently occuring letter.


```python
def freq_letters(text):
    '''
    For textual input, returns the most frequently occuring letter 
    (your function should be case-insensitive)
    '''
    alpha=[]
    for i in text:
        if str.isalpha(i)==True:
            alpha.append(i)
            
    res={}
    for j in alpha:
        res[j]=alpha.count(j)
        
    for val in res.values():
        return max(res)
    ...
```

```python
grader.check("q3_b_1")
```

Output:
```
q3_b_1 results: All test cases passed!
q3_b_1 - 1 message: Good job!
```

Write a second function that returns the number of vowels and consonants in the text.

```python
def vowels(text):
    '''
    For textual input, returns the number of vowels and consonants.
    Return a tuple with number of vowels as the first element and consonants
    as the second element, i.e (num_vowels, num_consonant)
    '''
    vowels=list('aeiouAEIOU')
    consonant=list('bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ')
    alpha=[]
    for i in text:
        if str.isalpha(i)==True:
            alpha.append(i)
            
    num_vowels=0
    num_consonant=0
    for j in alpha:
        if j in vowels:
            num_vowels+=1
        elif j in consonant:
            num_consonant+=1
    print ("(", num_vowels, ", ", num_consonant, ")", sep='')
    ...
```

```python
grader.check("q3_b_2")
```

Output:
```
q3_b_2 results:
    q3_b_2 - 1 result:
        Trying:
            vowels('I like banana!!') == (6, 5)
        Expecting:
            True
        **********************************************************************
        Line 1, in q3_b_2 0
        Failed example:
            vowels('I like banana!!') == (6, 5)
        Expected:
            True
        Got:
            (6, 5)
            False
```

## Part C

The question in this part of the assignment is intended to:

  1) Illuminate how what you already know is useful in this context.
  
  2) Provide an exercise in writing about coding concepts and ideas and using the relevant language and terminology. 
  
  3) Demonstrate your understanding of the assignment material in a different format.


<!-- BEGIN QUESTION -->

### Question 1 
(3 points) 

In 100 to 250 words identify a concept you have found difficult or confusing from this assignment. Reflect on how your previous learning or experience helped you to understand this concept. Provide your reflection using markdown in the cell below. 

The most difficult Part for me was Part B Question 3, because I was trapped by how to make function return the letter corresponding to the highest number of occurrences after counting the number of occurrences of each letter. So, I reviewed my notes from my undergraduate Python class and finally I found that Dictionary is used to store data in pairs like "key: value", then the last question was solved: I created an empty dictionary with each letter as a key, and each number of occurrences as a value, finally let function return the key corresponding to the largest value.

<!-- END QUESTION -->

## Submission

Make sure you have run all cells in your notebook and that your code works as expected. Save your notebook and upload it to <a href=https://www.gradescope.ca >Gradescope</a>
**Please save (ctrl+s) before uploading!**

```python

```

Tests for q1_a (points: 3)
```
>>> cent_binomial_coef(0) == 1
True
```
failure_message: Don't forget (0)! is 1.
```
>>> cent_binomial_coef(1) == 2
True
```


Tests for q1_b (points: 5)
```
>>> def testcase1(intersection_list):
...     temp = intersection_list([2,1,2,5,3,1], [2,1,7,1])
...     length = len(temp)
...     ones = temp.count(1)
...     twos = temp.count(2)
...     assert (length, ones, twos) == (3, 2, 1)
>>> testcase1(intersection_list)
```


Tests for q2_a (points: 8)
```
>>> is_prime(2) == True
True
```
```
>>> is_prime(3) == True
True
```
```
>>> is_prime(4) == False
True
```


Tests for q2_b (points: 8)
```
>>> striple_prod([5, 1, 2], [-1]*3, [2, 1, 0]) == 5
True
```


Tests for q3_a (points: 5)
```
>>> population(5.7) == 16
True
```
```
>>> population(2.9) == 12
True
```


Tests for q3_b_1 (points: 5)
```
>>> freq_letters('a!!!!!').lower() == 'a'
True
```
failure_message: You should only count letters a-z and not numbers, punctuation, spaces, etc.
success_message: Good job!


Tests for q3_b_2 (points: 5)
```
>>> vowels('I like banana!!') == (6, 5)
True
```
